# Programming in Python for Data Science 

# Assignment 4: Basic Python

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are the links of 2 videos that explain:

1. [How to save your Jupyter notebook](https://youtu.be/0aoLgBoAUSA) and,       
2. [How to answer a question in a Jupyter notebook assignment](https://youtu.be/7j0WKhI3W4s).    

### Assignment Learning Goals:

By the end of the module, students are expected to:

- Compare and contrast Python's key data types.
- Compare and contrast Python's key data structures. 
- Use python to determine the type and structure of an object.
- Demonstrate how to create data structures and convert them to another.
- Identify which operations can be applied to different data types and columns dtypes. 

This assignment covers [Module 4](https://prog-learn.mds.ubc.ca/en/module4) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` and the `raise NotImplementedError # No Answer - remove if you provide an answer` with your completed code and answers then proceed to run the cell!

Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct. These questions are worth 2 points.

```python
# Import libraries needed for this lab
import pandas as pd
import test_assignment4 as t
from hashlib import sha1
import altair as alt
```

## 1.  Python Datatypes 

**Question 1(a)** <br> {points: 1}  

Match the objects below to their associated correct datatype.
 
    a) 'False'     
    b) 67       
    c) None      
    d) [False, 'Wrong', None, "Lies"]      
    e) True      
    f) {'Name': 'Cookies and Cream', 'Brand': "Hersey's", 'Weight' = 48}        
    g) "None"      
    h) 67.0        
    i) ('Never', 'gonna', 'give', 'you', 'up')          
    j) 7.777777          
    k) "45"          


In the code chunk below, assign each object to it's corresponding datatype by putting the  following datatypes in quotes (as a string!):       
`int`, `float`, `bool`, `str`, `list`, `tuple`, `dict`, `NoneType`.


_Make sure your answer is written as the datatype selected above. _

***Example:***

    z)   23

`answer1_z = 'int'`

<!-- grade_id: cell-dd23213c014f85af, points: 0, grade: false, solution: true -->
```python
answer1_a = 'str'
answer1_b = 'int'
answer1_c = 'NoneType'
answer1_d = 'list'
answer1_e = 'bool'
answer1_f = 'dict'
answer1_g = 'str'
answer1_h = 'float'
answer1_i = 'tuple'
answer1_j = 'float'
answer1_k = 'str'
```

<!-- grade_id: cell-e7ceb2b5a029de7b, points: 1, grade: true, solution: false -->
```python
t.test_1a([answer1_a,answer1_b,answer1_c,answer1_d,answer1_e,answer1_f,answer1_g,answer1_h,
         answer1_i,answer1_j,answer1_k])
```

Output:
```
'Success'
```

## 2. Pandas Datatypes 

For the next couple questions, please refer to this table referencing different astronauts and their space accomplishments: 

|    | Name |   Year |   Group | Status   | Birth_Date | Birth_Place    | Gender   | Graduate | Military | Space_Walks | Space_Walks_hrs | Deceased   | Missions  |
|---:|:-----|----:|-----:|:-----|:-------|:--------|:---------|:---------|:------|-----:|-----:|:-----------|:-----|
|**0**| Neil A. Armstrong |   1962 |2 | Deceased | 1930-08-05| Wapakoneta, OH | Male| Aerospace Engineering   | False | 1 |  2   | True  | [Gemini 8,  Apollo 11]   |
|**1**| Buzz Aldrin  |1963 |  3 | Retired  | 1930-01-20| Montclair, NJ  | Male     | Astronautics | True | 2 | 8 | False      | [ Gemini 12,  Apollo 11] |
|**2**| Patricia Hilliard Robertson |   1998 |   17 | Deceased | 1963-03-12| Indiana, PA | Female   | Medicine| False| 0 |0   | True| [] |
|**3**| Alfred M. Worden |   1966 |  5 | Retired  | 1932-02-07| Jackson, MI    | Male |Aeronautical & Astronautical Engineering | True | 1 | 0.5 | False      | [Apollo 15]|


**Question 2(a)** <br> {points: 1}  

For the dataframe above what is the data dtype of each column? Save your answer as a string in an object named with the column label.      
Remember you can chose from the following data types: 

- `object` 
- `int64`   
- `float64`   
- `datetime64`   
- `bool`     

<!-- grade_id: cell-550e922f66bcbd57, points: 0, grade: false, solution: true -->
```python
Name = 'object'
Year = 'int64'
Group = 'int64'
Status = 'object'
Birth_Date = 'datetime64'
Birth_Place = 'object'
Gender = 'object'
Graduate = 'object'
Military = 'bool'
Space_Walks = 'int64'
Space_Walks_hrs = 'float64'
Deceased = 'bool'
Missions = 'object'
```

<!-- grade_id: cell-6ad4da99b7dfef74, points: 1, grade: true, solution: false -->
```python
t.test_2a([
    Name,
    Year,
    Group,
    Status,
    Birth_Date,
    Birth_Place,
    Gender,
    Graduate,
    Military,
    Space_Walks,
    Space_Walks_hrs,
    Deceased,
    Missions
])
```

Output:
```
'Success'
```

**Question 2(b)** <br> {points: 2} 

Using the dataframe above, which of the columns contain cell values that are **NOT** of type `str`?

_Assign your answer to an object called `answer2_b` in the code chunk below. Place each column name in a list._

_For example if you believe the answer is `Name`, `Year` and `Group` your answer would like this:_  

`answer2_b = ['Name', 'Year', 'Group']`


<!-- grade_id: cell-d1ca33359d1a6e78, points: 0, grade: false, solution: true -->
```python
answer2_b = [
    'Year',
    'Group',
    'Birth_Date',
    'Military',
    'Space_Walks',
    'Space_Walks_hrs',
    'Deceased',
    'Missions'
]
```

<!-- grade_id: cell-8a2c66ce17162ab8, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer2_b' in globals(
), "Please make sure that your solution is named 'answer2_b'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 2(c)** <br> {points: 1}  

Let's try the same time with a new dataframe below: 

|  | show_id| categ|title|director| cast|us_made|date_added|release_year| rating| duration| listed_in| description|
|--:|:-------|:----|:-----|:--------|:----|:------|:---------|-----------:|:------|:--------|:---------|------------:|
|**0**| 70044595  | Movie   | Monster House   | Gil Kenan   | [Mitchel Musso, Sam Lerner, Spencer Locke, Steve Buscemi, Kathleen Turner, Jason Lee, Jon Heder, Maggie Gyllenhaal] | TRue  | 01-Nov-19  | 2006 | PG  | 91 min | Children & Family Movies, Comedies   |  NaN |
|**1**| 7oo56441  | Movie   | Teenage Mutant Ninja Turtles| Kevin Munroe    | [Chris Evans, Sarah Michelle Gellar, Mako Iwamatsu, Kevin Smith, Patrick Stewart, Zhang Ziyi, Laurence Fishburne]   | NaN| 01-May-19|2007 | PG| 87 min     | Children & Family Movies, Comedies   |NaN |
|**2**| 80164032  | Movie   | A Plastic Ocean| NaN | Tanya Streeter | False| 04-Oct-19|2016 | PG       | 100 min    | Documentaries  |NaN |
|**3**| 80036412  | TV Show | A.D. Kingdom and Empire  | NaN | [Juan Pablo Di Pace, Adam Levy, Chipo Chung, Babou Ceesay, Emmett Scanlan, Will Thorp, Richard Coyle] | True.     | 15-Dec-17|2015 | 14| 46 mins| TV Dramas|NaN |
|**4**| 70196145  | TV Show | New Girl| NaN| [Zooey Deschanel, Jake Johnson, Max Greenfield, Hannah Simone, Lamorne Morris, Damon Wayans Jr.]| True| April 11, 2019 |2017 | 14| 28 mins  | Romantic TV Shows, TV Comedies|NaN |
|**5**| 60029154  | Movie| Terminator 3: Rise of the Machines | Jonathan Mostow | [Arnold Schwarzenegger, Nick Stahl, Claire Danes, Kristanna Loken, David Andrews, Mark Famiglietti]| True| 01-Jan-20|2003 | R| 109 min    | Action & Adventure, Sci-Fi & Fantasy |NaN |

What are the data types of the values within each column?

Save your answers as a string in an object named with the column label.      
Remember you can chose from the following data types: 

- `str` 
- `int`   
- `float`   
- `Timestamp`   
- `bool`   


<!-- grade_id: cell-e9ff05447014e52d, points: 0, grade: false, solution: true -->
```python
show_id = 'str'
categ = 'str'
title = 'str'
director = 'str'
cast = 'str'
us_made = 'str'
date_added = 'str'
release_year = 'int'
rating = 'str'
duration = 'str'
listed_in = 'str'
description = 'float'
```

<!-- grade_id: cell-499dd84702e1817a, points: 1, grade: true, solution: false -->
```python
t.test_2c([
    show_id,
    categ,
    title,
    director,
    cast,
    us_made,
    date_added,
    release_year,
    rating,
    duration,
    listed_in,
    description
])
```

Output:
```
'Success'
```

# 3. Gapminder 

Remember the `gapminder.csv` data that we used in Assignment 3. 

```python
gapminder = pd.read_csv('data/gapminder.csv')
display(gapminder)
```

Output:
```
Country      2012      2013      2014      2015      2016
0    Afghanistan  31200000  32300000  33400000  34400000  35400000
1        Albania   2910000   2900000   2900000   2890000   2890000
2        Algeria  37400000  38100000  38900000  39700000  40600000
3        Andorra     82400     80800     79200     78000     77300
4         Angola  25100000  26000000  26900000  27900000  28800000
..           ...       ...       ...       ...       ...       ...
190    Venezuela  29400000  29800000  30000000  30100000  29900000
191      Vietnam  89800000  90800000  91700000  92700000  93600000
192        Yemen  24500000  25100000  25800000  26500000  27200000
193       Zambia  14500000  14900000  15400000  15900000  16400000
194     Zimbabwe  13100000  13400000  13600000  13800000  14000000

[195 rows x 6 columns]
```

**Question 3(a)** <br> {points: 1} 

Use `.dtypes` to find the data type of each column.

<!-- grade_id: cell-43a3fe0102f54fe9, points: 0, grade: false, solution: true -->
```python
country = gapminder.dtypes['Country']
year_2012 = gapminder.dtypes['2012']
year_2013 = gapminder.dtypes['2013']
year_2014 = gapminder.dtypes['2014']
year_2015 = gapminder.dtypes['2015']
year_2016 = gapminder.dtypes['2016']

print(f'Country : {country}')
print(f'2012 : {year_2012}')
print(f'2013 : {year_2013}')
print(f'2014 : {year_2014}')
print(f'2015 : {year_2015}')
print(f'2016 : {year_2016}')
```

Output:
```
Country : object
2012 : int64
2013 : int64
2014 : int64
2015 : int64
2016 : int64
```

<!-- grade_id: cell-d5755344476551c4, points: 1, grade: true, solution: false -->
```python
t.test_3a([country,year_2012,year_2013,year_2014,year_2015,year_2016])
```

Output:
```
'Success'
```

**Question 3(b)** <br> {points: 1} 

Similarly as we did in the last assignment, melt the data to produce a single column named `Year` that contain the years 2012 - 2014 and a new column named `Population` with associated population values. Name the long dataframe `gapminder_long`. 


<!-- grade_id: cell-b9790e4a26b35d2c, points: 0, grade: false, solution: true -->
```python
gapminder_long = gapminder.melt(
    id_vars = ['Country'],
    value_vars = [str(year) for year in range(2012, 2014 + 1)],
    var_name = 'Year',
    value_name = 'Population'
).sort_values(['Year', 'Population'], ascending = True)
display(gapminder_long)
```

Output:
```
Country  Year  Population
72        Holy See  2012         786
119          Nauru  2012       10100
181         Tuvalu  2012       10700
131          Palau  2012       17600
145     San Marino  2012       32100
..             ...   ...         ...
413         Brazil  2014   203000000
467      Indonesia  2014   255000000
576  United States  2014   319000000
466          India  2014  1300000000
425          China  2014  1400000000

[585 rows x 3 columns]
```

<!-- grade_id: cell-09e08d60d57a2a2a, points: 1, grade: true, solution: false -->
```python
t.test_3b(gapminder_long)
```

Output:
```
'Success'
```

**Question 3(c)** <br> {points: 1} 


What are the datatypes of the columns `Year` and `Population` in the `gapminder_long` dataframe?   
 Save them in the objects named `Year_type` and `Population_type`.

<!-- grade_id: cell-68ddafc972be022a, points: 0, grade: false, solution: true -->
```python
Year_type = 'object'
Population_type = 'int64'
```

<!-- grade_id: cell-a1d43232711c4f89, points: 1, grade: true, solution: false -->
```python
t.test_3c([Year_type,Population_type])
```

Output:
```
'Success'
```

**Question 3(d)** <br> {points: 2} 

If we attempt to take the range of the years in the data we obtain the following error: 

```python
# Uncomment the code below to see the error
# max(gapminder_long['Year']) - min(gapminder_long['Year'])
```

Output:
```
TypeError                                 Traceback (most recent call last)
c:\My_Files\School\School_Work\Engineering_Bachelors\2022\KeyCapabilities_DataScience\prog-python-ds-students\release\assignment4\assignment4.ipynb Cell 35 in <cell line: 2>()
      <a href='vscode-notebook-cell:/c%3A/My_Files/School/School_Work/Engineering_Bachelors/2022/KeyCapabilities_DataScience/prog-python-ds-students/release/assignment4/assignment4.ipynb#X46sZmlsZQ%3D%3D?line=0'>1</a> # Uncomment the code below to see the error
----> <a href='vscode-notebook-cell:/c%3A/My_Files/School/School_Work/Engineering_Bachelors/2022/KeyCapabilities_DataScience/prog-python-ds-students/release/assignment4/assignment4.ipynb#X46sZmlsZQ%3D%3D?line=1'>2</a> max(gapminder_long['Year']) - min(gapminder_long['Year'])

TypeError: unsupported operand type(s) for -: 'str' and 'str'
```

Which of the following most appropriately explains what is happening with the `Year` variable. 


A) When `.melt()` was used, it converted the `Year` values from data types `int` to `str`.  

B) The `Year` column labels were always of type `str`. When melted, it converted the labels to cell values so the labels became cell values of type `str` and we cannot use a subtraction operator on values of type `str`. 

C) The `Year` column labels were always of type `str`. When melted, it converted the labels to cell values so the labels became cell values of type `str`  and we cannot use the verbs `max()` and `min()` on values of type `str`. 

D) There is no column named `Year` and so the code cannot execute. 

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer3_d`.*

<!-- grade_id: cell-d7eb6e529f418cad, points: 0, grade: false, solution: true -->
```python
answer3_d = 'B'
```

<!-- grade_id: cell-b356d614209a82a2, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer3_d' in globals(
), "Please make sure that your solution is named 'answer3_d'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 3(e)** <br> {points: 1} 


Tranform the column `Year` in the `gapminder_long` dataframe into a more appropriate dtype.     
Name the new dataframe object `gm_long_fixed`.

<!-- grade_id: cell-10b02c1e5055cd3a, points: 0, grade: false, solution: true -->
```python
gm_long_fixed = gapminder_long.assign(Year = lambda x : x['Year'].astype('int64'))
display(gm_long_fixed)
```

Output:
```
Country  Year  Population
72        Holy See  2012         786
119          Nauru  2012       10100
181         Tuvalu  2012       10700
131          Palau  2012       17600
145     San Marino  2012       32100
..             ...   ...         ...
413         Brazil  2014   203000000
467      Indonesia  2014   255000000
576  United States  2014   319000000
466          India  2014  1300000000
425          China  2014  1400000000

[585 rows x 3 columns]
```

<!-- grade_id: cell-15fc50ced0793e77, points: 1, grade: true, solution: false -->
```python
t.test_3e(gm_long_fixed)
```

Output:
```
'Success'
```

**Question 3(f)** <br> {points: 1} 

After casting the column `Year` into a more appropriate data type, does the code for obtaining the range of years in the data (from question 3d) execute without errors with `gm_long_fixed`? 

A) No, since the verbs `max()` and `min()` cannot be used on values with the new data type.

B) No, since subtraction cannot be done on on values with the new data type.

C) Yes it executes and the range is 2.

D) Yes, it executes and the range is 1. 

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer3_f`.*

```python
max(gm_long_fixed['Year']) - min(gm_long_fixed['Year'])
```

Output:
```
2
```

<!-- grade_id: cell-073b3d477d49274b, points: 0, grade: false, solution: true -->
```python
answer3_f = 'C'
```

<!-- grade_id: cell-4b62ed87febc005d, points: 1, grade: true, solution: false -->
```python
t.test_3f(answer3_f)
```

Output:
```
'Success'
```

## 4. Lists, Sets and Tuples 

**Question 4** <br> {points: 10}  

Which data structure(s) describe the following statements  (multiple datatypes are possible): 

    a) Elements cannot be accessed with square brackets 
    b) Can be sliced, take repeated values, and entries can be changed and edited.   
    c) Cannot be sliced and entries can be independently deleted.   
    d) A set of key value pairs.  
    e) Immutable.  
    f) Are created with parentheses.   
    g) Can be sliced.   
    h) Can be split into their separate entries with `.split()`.  
    i) Outputted when using `.split()`.  
    j) Are created with curly brackets.    

 <br> 


Save all possible answers as strings within a list.      
Remember you can chose from the following data types: 

- `str` 
- `list`   
- `tuple`   
- `set`   
- `dict`   

***Example:***    

`answer4_z = ['set', 'dict']`



<!-- grade_id: cell-9d07992c6bd83587, points: 0, grade: false, solution: true -->
```python
answer4_a = ['set']
answer4_b = ['list']
answer4_c = ['set', 'dict']
answer4_d = ['dict']
answer4_e = ['str', 'tuple']
answer4_f = ['tuple']
answer4_g = ['str', 'list', 'tuple']
answer4_h = ['str']
answer4_i = ['list']
answer4_j = ['set', 'dict']
```

<!-- grade_id: cell-12bb2d09caa43239, points: 10, grade: true, solution: false -->
```python
t.test_4a([
    sorted(answer4_a), 
    sorted(answer4_b),
    sorted(answer4_c),
    sorted(answer4_d),
    sorted(answer4_e),
    sorted(answer4_f),
    sorted(answer4_g),
    sorted(answer4_h),
    sorted(answer4_i),
    sorted(answer4_j)
])
```

Output:
```
'Success'
```

## 5. Workforce Pay Rates

The data named `workforce-pay-rates-and-gender.csv`  was aquired from the [City of Vancouver's Open Data Portal](https://opendata.vancouver.ca/explore/dataset/workforce-pay-rates-and-gender/information/?sort=classification). 
It provides a look at the range of hourly wages for city jobs across Vancouver including the gender of staff within sector classification.  

***Note: many popular datasets have sex as a feature where the possible values are male and female. This representation reflects how the data were collected and is not meant to imply that, for example, gender is binary.***



The data we are reading in is contained in a json file. It looks like the output of the object `data` below.

```python
# ignore this code and instead focus on the output of data in the next cell
import json

with open('data/workforce-pay-rates-and-gender.json') as file:
    data = json.load(file)
```

```python
# Uncomment out the next line to view the object

display(data)
```

Output:
```
{'Year': {'0': 2019,
  '1': 2019,
  '2': 2019,
  '3': 2019,
  '4': 2019,
  '5': 2019,
  '6': 2019,
  '7': 2019,
  '8': 2019,
  '9': 2019,
  '10': 2019,
  '11': 2019,
  '12': 2019,
  '13': 2019,
  '14': 2019,
  '15': 2019,
  '16': 2019,
  '17': 2019,
  '18': 2019,
  '19': 2019,
  '20': 2019,
  '21': 2019,
  '22': 2019,
  '23': 2019,
  '24': 2019,
  '25': 2019,
  '26': 2019,
  '27': 2019,
  '28': 2019,
  '29': 2019,
  '30': 2019,
  '31': 2019,
  '32': 2019,
  '33': 2019,
  '34': 2019,
  '35': 2019,
  '36': 2019,
  '37': 2019,
  '38': 2019,
  '39': 2019,
  '40': 2019,
  '41': 2019,
  '42': 2019,
  '43': 2019,
  '44': 2019,
  '45': 2019,
  '46': 2019,
  '47': 2019,
  '48': 2019,
  '49': 2019,
  '50': 2019,
  '51': 2019,
  '52': 2019,
  '53': 2019,
  '54': 2019,
  '55': 2019,
  '56': 2019,
  '57': 2019,
  '58': 2019,
  '59': 2019,
  '60': 2019,
  '61': 2019,
  '62': 2019,
  '63': 2019,
  '64': 2019,
  '65': 2019,
  '66': 2019,
  '67': 2019,
  '68': 2019,
  '69': 2019,
  '70': 2019,
  '71': 
```

**Question 5(a)**  <br> {points: 1}  

What kind of data structure is a json file? 

A) List 

B) Tuple 

C) Dictionary

D) Set

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer5_a`.*

<!-- grade_id: cell-b5c598b1591c8a7b, points: 0, grade: false, solution: true -->
```python
answer5_a = 'C'
```

<!-- grade_id: cell-870615ffdfe8cbaf, points: 1, grade: true, solution: false -->
```python
t.test_5a(answer5_a)
```

Output:
```
'Success'
```

**Question 5(b)**  <br> {points: 1}  

Us the appropriate verb to convert data into a dataframe.    
Name the object `van_wages`.

<!-- grade_id: cell-d467a43d7301b394, points: 0, grade: false, solution: true -->
```python
van_wages = pd.DataFrame.from_dict(data)

display(van_wages)
```

Output:
```
Year            Classification  Union  Minimum hourly rate  \
0    2019  Administrative Manager 3  False                47.97   
1    2019         General Manager 2  False               120.47   
2    2019                 Manager 1  False                50.06   
3    2019                 Manager 3  False                64.62   
4    2019      Operations Manager 0  False                40.62   
..    ...                       ...    ...                  ...   
184  2019             Lead Mechanic   True                53.67   
185  2019     Training Officer Fire   True                71.43   
186  2019          Theatre Casual V   True                32.87   
187  2019        Theatre Technician   True                40.35   
188  2019     Electrical Apprentice   True                17.49   

     Maximum hourly rate  Women   Men  
0                  66.30   18.0  19.0  
1                 150.95    3.0   6.0  
2                  70.97   49.0  39.0  
3                  95.49   19.0  32.0  

```

<!-- grade_id: cell-83f264ae29621331, points: 1, grade: true, solution: false -->
```python
t.test_5b(van_wages)
```

Output:
```
'Success'
```

**Question 5(c)**  <br> {points: 1}  


Created a column for the total number of people with each classification, and name the column  `Total`. 
Name the new object `van_wages_total`.


```python
van_wages.sort_values(['Year', 'Classification'], ascending = True, inplace = True)
display(van_wages)
```

Output:
```
Year           Classification  Union  Minimum hourly rate  \
148  2019       Admin Management I   True                33.60   
151  2019     Admin Management III   True                46.75   
152  2019    Admin Management IIIa   True                39.56   
150  2019     Admin Management IIa   True                34.98   
153  2019      Admin Management IV   True                43.04   
..    ...                      ...    ...                  ...   
27   2019    Working Supervisor II   True                32.40   
92   2019   Working Supervisor III   True                34.68   
93   2019    Working Supervisor IV   True                35.75   
145  2019    Working Supervisor IX   True                41.22   
147  2019  Working Supervisor VIII   True                39.16   

     Maximum hourly rate  Women   Men  
148                37.97    1.0   5.0  
151                46.75    1.0   3.0  
152                46.75   10.0   9.0  
150                43.04    2.0   NaN  
153         
```

```python
van_wages_total = van_wages.fillna(0).assign(Total = lambda x : x['Women'] + x['Men'])
display(van_wages_total)
```

Output:
```
Year           Classification  Union  Minimum hourly rate  \
148  2019       Admin Management I   True                33.60   
151  2019     Admin Management III   True                46.75   
152  2019    Admin Management IIIa   True                39.56   
150  2019     Admin Management IIa   True                34.98   
153  2019      Admin Management IV   True                43.04   
..    ...                      ...    ...                  ...   
27   2019    Working Supervisor II   True                32.40   
92   2019   Working Supervisor III   True                34.68   
93   2019    Working Supervisor IV   True                35.75   
145  2019    Working Supervisor IX   True                41.22   
147  2019  Working Supervisor VIII   True                39.16   

     Maximum hourly rate  Women   Men  Total  
148                37.97    1.0   5.0    6.0  
151                46.75    1.0   3.0    4.0  
152                46.75   10.0   9.0   19.0  
150                43.04
```

<!-- grade_id: cell-de942db5679929f8, points: 1, grade: true, solution: false -->
```python
t.test_5c(van_wages_total)
```

Output:
```
'Success'
```

**Question 5(d)**  <br> {points: 2}  

Which columns in `van_wages_total` will throw an error if the ***sum*** of the column is attempted to be calculated.  

Save the column name(s) as strings in a list as an object named `answer5_d`.        

***Example:***

`answer6_e = ['Year', 'Union', 'Minimum hourly rate']`


```python
van_wages_total.dtypes
```

Output:
```
Year                     int64
Classification          object
Union                     bool
Minimum hourly rate    float64
Maximum hourly rate    float64
Women                  float64
Men                    float64
Total                  float64
dtype: object
```

<!-- grade_id: cell-0de07d28b19d1011, points: 0, grade: false, solution: true -->
```python
answer5_d = ['Classification']
```

<!-- grade_id: cell-db4a5922a7b5d25a, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer5_d' in globals(
), "Please make sure that your solution is named 'answer5_d'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.

```

**Question 5(e)**  <br> {points: 1}  

Which columns in `van_wages_total` will throw an error if the ***mean*** of the column is attempted to be calculated.

Save the column name(s) as strings in a list as an object named `answer5_e`.    

***Example:***

`answer5_e = ['Year', 'Union', 'Minimum hourly rate']`


If you think that none of the columns will throw an error than answer will be an empty list `[]`.  

<!-- grade_id: cell-e4dbbeb6eb1acd64, points: 0, grade: false, solution: true -->
```python
answer5_e = ['Classification']
```

<!-- grade_id: cell-cd94178956ccec26, points: 1, grade: true, solution: false -->
```python
t.test_5e(answer5_e)
```

Output:
```
'Success'
```

## 6. Lists and Operations with Weather

Below is the upcoming 7 day weather forecast  displayed in a dataframe: 

|    | day       | description       |   day_temperature |   night_temperature |   pop | wind   |   hrs_sun | rain   |
|---:|:----------|:------------------|------------------:|--------------------:|------:|:-------|----------:|:-------|
|  0 | Monday    | sunny             |                22 |                  14 |    10 | 7 SW   |        15 |        |
|  1 | Tuesday   | sunny             |                21 |                  15 |     0 | 8 S    |        15 |        |
|  2 | Wednesday | showers           |                23 |                  14 |    80 | 13 SW  |         6 | 1-3 mm |
|  3 | Thursday  | cloudy            |                23 |                  14 |    20 | 10 W   |         9 |        |
|  4 | Friday    | chance of showers |                22 |                  14 |    70 | 8 SE   |         5 | ~1 mm  |
|  5 | Saturday  | chance of showers |                21 |                  12 |    40 | 21 W   |         4 | ~1 mm  |
|  6 | Sunday    | chance of showers |                22 |                  13 |    40 | 12 W   |         5 | ~1 mm  |

**Question 6(a)**  <br> {points: 1}

Create the data frame using the data structures provided below and name the dataframe `weather`. 

```python
monday = ['Monday', 'sunny', 22.0, '14.0', '10', '7 SW', 15.0, None]
tuesday = ['Tuesday', 'sunny', 21.0, '15.0', '0', '8 S', 15.0, None]
wednesday = ['Wednesday', 'showers', 23.0, '14.0', '80', '13 SW', 6.0, '1-3 mm']
thursday = ['Thursday', 'cloudy', 23.0, '14.0', '20', '10 W', 9.0, None]
friday = ['Friday', 'chance of showers', 22.0, '14.0', '70', '8 SE', 5.0, '~1 mm']
saturday = ['Saturday', 'chance of showers', 21.0, '12.0', '40', '21 W', 4.0, '~1 mm']
sunday = ['Sunday', 'chance of showers', 22.0, '13.0', '40', '12 W', 5.0, '~1 mm']


names = ['day', 'description',  'day_temperature', 'night_temperature', 'pop', 'wind', 'hrs_sun', 'rain']
```

<!-- grade_id: cell-cd6d1b9a30fd24c8, points: 0, grade: false, solution: true -->
```python
data = [
    monday, 
    tuesday, 
    wednesday, 
    thursday, 
    friday, 
    saturday, 
    sunday, 
]

weather = pd.DataFrame(data, columns = names)
```

Output:
```
[['Monday', 'sunny', 22.0, '14.0', '10', '7 SW', 15.0, None],
 ['Tuesday', 'sunny', 21.0, '15.0', '0', '8 S', 15.0, None],
 ['Wednesday', 'showers', 23.0, '14.0', '80', '13 SW', 6.0, '1-3 mm'],
 ['Thursday', 'cloudy', 23.0, '14.0', '20', '10 W', 9.0, None],
 ['Friday', 'chance of showers', 22.0, '14.0', '70', '8 SE', 5.0, '~1 mm'],
 ['Saturday', 'chance of showers', 21.0, '12.0', '40', '21 W', 4.0, '~1 mm'],
 ['Sunday', 'chance of showers', 22.0, '13.0', '40', '12 W', 5.0, '~1 mm']]
         day        description  day_temperature night_temperature pop   wind  \
0     Monday              sunny             22.0              14.0  10   7 SW   
1    Tuesday              sunny             21.0              15.0   0    8 S   
2  Wednesday            showers             23.0              14.0  80  13 SW   
3   Thursday             cloudy             23.0              14.0  20   10 W   
4     Friday  chance of showers             22.0              14.0  70   8 SE   
5   Saturday  chance of sh
```

<!-- grade_id: cell-b06c0f51a6768b9f, points: 1, grade: true, solution: false -->
```python
t.test_6a(weather)
```

Output:
```
'Success'
```

**Question 6(b)**  <br> {points: 2}

What data types or data structure are the object labeled by the days of the week? 


A) List 

B) Tuple 

C) Dictionary

D) Set

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer6_b`.*

```python
weather.dtypes
```

Output:
```
day                   object
description           object
day_temperature      float64
night_temperature     object
pop                   object
wind                  object
hrs_sun              float64
rain                  object
dtype: object
```

<!-- grade_id: cell-906749f0f0237903, points: 0, grade: false, solution: true -->
```python
answer6_b = 'str'
```

<!-- grade_id: cell-c1fa535dc6b05f3a, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer6_b' in globals(
), "Please make sure that your solution is named 'answer6_b'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 6(c)**  <br> {points: 1}

What is the average chance of precipitation (`pop`) over the week? Answer using the `.mean()` function and convert datatypes if needed.     
Save as object as `pop_mean`.

<!-- grade_id: cell-1612bf4df887bd04, points: 0, grade: false, solution: true -->
```python
pop_mean = weather['pop'].astype('float64').mean()
pop_mean
```

Output:
```
37.142857142857146
```

<!-- grade_id: cell-ca0adf0f075431e0, points: 1, grade: true, solution: false -->
```python
t.test_6c(pop_mean)
```

Output:
```
'Success'
```

**Question 6(d)**  <br> {points: 1}

Create a new column named `temp_avg` which is the average temperature for the whole day (day and night temperature).    
Name the new dataframe `weather_ext`.       
_Hint: You may need to make changes to the `night_temperature` column first._

<!-- grade_id: cell-75bffa5b60a04627, points: 0, grade: false, solution: true -->
```python
weather_ext = weather.assign(night_temperature = lambda x : x['night_temperature'].astype('float64'))
weather_ext = weather_ext.assign(temp_avg = lambda x : (x['day_temperature'] + x['night_temperature'])/2)
display(weather_ext)
```

Output:
```
day        description  day_temperature  night_temperature pop  \
0     Monday              sunny             22.0               14.0  10   
1    Tuesday              sunny             21.0               15.0   0   
2  Wednesday            showers             23.0               14.0  80   
3   Thursday             cloudy             23.0               14.0  20   
4     Friday  chance of showers             22.0               14.0  70   
5   Saturday  chance of showers             21.0               12.0  40   
6     Sunday  chance of showers             22.0               13.0  40   

    wind  hrs_sun    rain  temp_avg  
0   7 SW     15.0    None      18.0  
1    8 S     15.0    None      18.0  
2  13 SW      6.0  1-3 mm      18.5  
3   10 W      9.0    None      18.5  
4   8 SE      5.0   ~1 mm      18.0  
5   21 W      4.0   ~1 mm      16.5  
6   12 W      5.0   ~1 mm      17.5
```

<!-- grade_id: cell-64b898d4611ff820, points: 1, grade: true, solution: false -->
```python
t.test_6d(weather_ext)
```

Output:
```
'Success'
```

**Question 6(e)**  <br> {points: 2}

Use `.str.split()` with argument `expand=True` to split the column `wind` into a 2 separate columns named  `wind_speed` and `wind_direction`. Make sure the `wind_speed` column is of dtype `int64`.    
Name the new dataframe `wind_stats`. 

<!-- grade_id: cell-95558162c7b975fc, points: 0, grade: false, solution: true -->
```python
wind_stats = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-5d8dc27af23d5065, points: 2, grade: true, solution: false -->
```python
t.test_6e(wind_stats)
```

**Question 6(f)**  <br> {points: 2}

Merge or concatenate these columns back with the original dataframe `weather_ext`. 
Name the new dataframe `weather_extended`.

<!-- grade_id: cell-6d5ae7710bf9d74d, points: 0, grade: false, solution: true -->
```python
weather_extended = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-5a4d7ddb2a290d79, points: 2, grade: true, solution: false -->
```python
t.test_6f(weather_extended)
```

**Question 6(g)**  <br> {points: 1}

Create a line plot using `alt.chart(..., width=500, height=300).mark_line()` where the x and y axis are `day` and `wind_speed` respectively. Don't forget to give it a title using the `.properties(title="")` function. 

If you want the days of the week to be sorted correctly, use the argument      
`sort=['Monday','Tuesday','Wednesday','Thursday','Friday','Saturday','Sunday']` within `alt.X()`. 

Name the plot `wind_plot`. 

<!-- grade_id: cell-812ff6784b477b3f, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-af571615e5396f56, points: 1, grade: true, solution: false -->
```python
t.test_6g(wind_plot)
```

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.  

## Attributions
- Gapminder Dataset - [Gapminder](https://www.gapminder.org/data/)
- Astronaut Dataset - [Kaggle](https://www.kaggle.com/nasa/astronaut-yearbook?select=astronauts.csv)
- Netflix Dataset - [Kaggle](https://www.kaggle.com/shivamb/netflix-shows?select=netflix_titles.csv)
- MDS DSCI 511 - Programming for Data Science - [MDS's GitHub website](https://ubc-mds.github.io/course-descriptions/DSCI_511_prog-dsci/) 


## Module Debriefing

If this video is not showing up below, click on the cell and click the ▶ button in the toolbar above.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('bHmaiwexLTc', width=854, height=480)
```

```python

```